# DS-GA-1007 Programming for Data Science

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Any *textual answers* that need to be provided will be marked with "YOUR ANSWER HERE". Replace this text with your answer to the question.

Any *code answers* that need to be provided will be marked with:

```
# YOUR CODE HERE
raise NotImplementedError()
```

Replace all this code with your answer to the question. If you do not answer the question, the `NotImplementedError` exception will be raised, which will indicate to the grader that no answer has been supplied.

In many cases, code answers will also have some associated test code. You should execute the tests after you have entered your code in order to ensure that your answer is correct. You should not proceed to the next question until your answer is correct.

Finally, insert your Net ID and the Net ID's of any collaborators in the cell below.

```python
NET_ID = "LH1036"
COLLABORATORS = ""
```

---

# Assignment 1

## 1. Function return value

This exercise is to show you how function return values work. Run the code below to see what happens.

```python
def f1(x):
    print(x + 1)

def f2(x):
    return(x + 1)

f1(4)
f2(4)
```

Output:
```
5
5
```

Each function should produce the same output. Now run the following code.

```python
f1(4) + 1
f2(4) + 1
```

Output:
```
5
TypeError                                 Traceback (most recent call last)
<ipython-input-3-02362f8b3abb> in <module>()
----> 1 f1(4) + 1
      2 f2(4) + 1

TypeError: unsupported operand type(s) for +: 'NoneType' and 'int'
```

Explain what is happening and why you get this result.

<!-- grade_id: Q1B, points: 1, grade: true, solution: true -->
f1 runs properly because it returns a value (4) that is an int. Addition of two integers is supported.
f2 throws a Type Error because it returns a None, and None and int "1" cannot be added/concatenated with "+".

## 2. Function operation

Consider the following piece of code:

```python
def f(x):
    if x == 0 or x == 1:
        return x
    return f(x - 1) + f(x - 2)
```

Describe, in words, what this code does, and how it does it.

<!-- grade_id: Q2A, points: 1, grade: true, solution: true -->
The function takes one argument (takes the place of parameter x). If the argument is equal to 0 or 1, the function returns the argument (i.e. 0 or 1). If the function is passed an argument that is an int or a floating point number in the form x.0 (e.g. 2.0, 2.00, 2.000) that is not equal to 0 or 1, the function is called with the argument minus 1 and the argument minus 2, and the sum is returned. This is repeated recursively for numbers > 2 until the recursion "stops" at 0 or 1.  E.g. f(3) returns f(2) + f(1); f(2) calls the function and returns f(1) + f(0), and f(1) returns 1. f(4) returns f(3) + f(2), where f(3) and f(2) return the values as just described.

For what inputs will this function not behave as expected? What will happen?

<!-- grade_id: Q2B, points: 1, grade: true, solution: true -->
The function does not behave as expected with inputs that are not integers or floating point numbers in the form x.0 (e.g. 1.0). If the input is a String or a list the function will throw a Type Error because these data types cannot accept the math operator for subtraction "-". If the input is a floating point number not in the form x.0 (e.g. "1.2" as opposed to "1.0") or a negative integer (e.g. -1), the function throws a Recursion Error because it can never "stop" at 0 or 1. If the function is passed too many arguments (even if they are ints), a Type Error is thrown. 

## 3. Defining a function

The following program plays the game of "Rock, Paper, and Scissors". Familiarize yourself with how it works. Try running the program.

```python
from random import randint
valid_input = ["R", "P", "S"]
winners = [["R", "S"], ["S", "P"], ["P", "R"]]

print("You are playing a Rock, Paper, Scissors simulation.")
user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
user_choice = user_choice.upper()

while user_choice != 'Q':
    if user_choice not in valid_input:
        print("Error: invalid input -- " + user_input)
        break
    
    computer_choice = valid_input[randint(0, len(valid_input)-1)]
    print("Computer chooses " + computer_choice)

    if user_choice == computer_choice:
        print("Tie Game")
    elif [user_choice, computer_choice] in winners:
        print("You win!")
    else:
        print("Computer wins!")
    
    user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
    user_choice = user_choice.upper()
```

Recall how we define a function using `def`, and how we pass in parameters. Transform the code into a function called `rps` that takes two parameters: `user_choice` and `computer_choice`, with valid inputs `'R'`, `'P'`, and `'S'`. Your function should return `'D'` for a draw, `'W'` if `user_choice` wins, and `'L'` if `user_choice` loses.

<!-- grade_id: Q3B, points: 0, grade: false, solution: true -->
```python
from random import randint
valid_input = ["R", "P", "S"]
winners = [["R", "S"], ["S", "P"], ["P", "R"]]
draws = [["R", "R"], ["S", "S"], ["P", "P"]]

def rps(user_input, computer_input):
    ''' function that return D, W, or L '''
    if [user_input, computer_input] in draws:
        return "D"
    elif [user_input, computer_input] in winners:
        return "W"
    else:
        return "L"

print("You are playing a Rock, Paper, Scissors simulation.")
user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
user_choice = user_choice.upper()

while user_choice != 'Q':
    if user_choice not in valid_input:
        print("Error: invalid input -- " + user_choice)
        break
    
    computer_choice = valid_input[randint(0, len(valid_input)-1)]
    print("Computer chooses " + computer_choice)
    
    result = rps(user_choice, computer_choice)
    if result == 'D':
        print("Tie Game")
    elif result == 'W':
        print("You win!")
    else:
        print("Computer wins!")
    
    user_choice = input("Select R for rock, P for Paper, S for scissors, Q to quit\n> ")
    user_choice = user_choice.upper()
```

Output:
```
You are playing a Rock, Paper, Scissors simulation.
Select R for rock, P for Paper, S for scissors, Q to quit
> r
Computer chooses S
You win!
Select R for rock, P for Paper, S for scissors, Q to quit
> p
Computer chooses S
Computer wins!
Select R for rock, P for Paper, S for scissors, Q to quit
> s
Computer chooses R
Computer wins!
Select R for rock, P for Paper, S for scissors, Q to quit
> s
Computer chooses R
Computer wins!
Select R for rock, P for Paper, S for scissors, Q to quit
> s
Computer chooses S
Tie Game
```

<!-- grade_id: Q3C, points: 1, grade: true, solution: false -->
from nose.tools import assert_equal
assert_equal(rps('R', 'R'), 'D')
assert_equal(rps('R', 'P'), 'L')
assert_equal(rps('R', 'S'), 'W')
assert_equal(rps('P', 'R'), 'W')
assert_equal(rps('P', 'P'), 'D')
assert_equal(rps('P', 'S'), 'L')
assert_equal(rps('S', 'R'), 'L')
assert_equal(rps('S', 'P'), 'W')
assert_equal(rps('S', 'S'), 'D')


## 4. List function

Given a list of non-empty tuples, return a list sorted in increasing order by the last element in each tuple. 

For example: `[(1, 7), (1, 3), (3, 4, 5), (2, 2)]` yields `[(2, 2), (1, 3), (3, 4, 5), (1, 7)]`

<!-- grade_id: Q4A, points: 0, grade: false, solution: true -->
```python
def sort_tuple_list(tuple_list):
    '''Function to sort a list of non-empty tuples in increasing order by the last element in each tuple'''
    
    if isinstance(tuple_list, list) and len(tuple_list) > 0:     # check that list is nonempty
        sorted_tuple = sorted(tuple_list, key = lambda tuple_list: tuple_list[-1]) # key = last element of each tuple
        return sorted_tuple
    
    else: # print warning if list is empty
        print("You did not give me a nonempty list!")
```

<!-- grade_id: Q4B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(sort_tuple_list([(1, 3), (3, 2), (2, 1)]), [(2, 1), (3, 2), (1, 3)])
assert_equal(sort_tuple_list([(2, 3), (1, 2), (3, 1)]), [(3, 1), (1, 2), (2, 3)])
assert_equal(sort_tuple_list([(1, 7), (1, 3), (3, 4, 5), (2, 2)]), [(2, 2), (1, 3), (3, 4, 5), (1, 7)])
```

# 5. List sort

Given two lists sorted in increasing order, create and return a merged list of all the elements in sorted order. You may modify the passed in lists. Ideally, the solution should work in "linear" time, making a single pass of both lists.

<!-- grade_id: Q5A, points: 0, grade: false, solution: true -->
```python
def linear_merge(list1, list2):
    '''Given two lists sorted in increasing order, create and 
       return a merged list of all the elements in sorted order
    '''
    concat_lists = list1 + list2 # merge lists into a new list
    return sorted(concat_lists) # return sorted list
```

<!-- grade_id: Q5B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(linear_merge(['aa', 'xx', 'zz'], ['bb', 'cc']), ['aa', 'bb', 'cc', 'xx', 'zz'])
assert_equal(linear_merge(['aa', 'xx'], ['bb', 'cc', 'zz']), ['aa', 'bb', 'cc', 'xx', 'zz'])
assert_equal(linear_merge(['aa', 'aa'], ['aa', 'bb', 'bb']), ['aa', 'aa', 'aa', 'bb', 'bb'])
```

# 6. Built-in Operator

Imagine that Python doesn’t have the `!=` operator built in. Write a function that takes two parameters and gives the same result as the `!=` operator. Obviously, you cannot use `!=` within your function! Test if your code works by thinking of examples and making sure the output is the same for your new method as `!=` gives you.

<!-- grade_id: Q6A, points: 0, grade: false, solution: true -->
```python
def not_equal(val1, val2):
    ''' User supplied version of the != operator, implemeted without using !='''
    return not (val1 == val2)
```

<!-- grade_id: Q6B, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_true, assert_false
assert_true(not_equal(2,3))
assert_false(not_equal(3,3))
assert_true(not_equal('3',3))
assert_true(not_equal((1,2),[1,2]))
assert_false(not_equal((3,4),(3,4)))
```

## 7. Quadratic Formula

A quadratic equation is a second order polynomial with three coefficients a, b, and c. It is given by:

$$ax^2 + bx + c = 0$$

The solutions to the equation are given by the formulas:

$$x_1 = \frac{-b + \sqrt{b^2 - 4ac}}{2a}$$

and

$$x_2 = \frac{-b - \sqrt{b^2 - 4ac}}{2a}$$



Write a function that computes the roots of a quadratic equation and returns them as a tuple. Your function should take three parameters, a, b, and c. Note that if

$$b^2 - 4ac < 0$$ 

then the roots will be complex. Hint: remember that Python has a built-in complex data type.

<!-- grade_id: Q7A, points: 0, grade: false, solution: true -->
```python
from math import sqrt

def roots(a, b, c):
    ''' Computes the roots of the quadratic equation with coefficients a, b, c, and returns
        them as the tuple (x1, x2)
    '''
    
    discriminant = (b*b) - (4*a*c) # check whether roots are real or imaginary
    
    if a != 0:

        if discriminant >= 0: # case 1: 2 real roots
            sqrt_dis = sqrt(discriminant)

            x1 = (-b + sqrt_dis)/(2*a)
            x2 = (-b - sqrt_dis)/(2*a)
            return (x1, x2)

        elif discriminant < 0: # case 2: 2 imaginary roots
            imaginary_dis = (sqrt(-1 * discriminant)) * 1j

            x1 = (-b + imaginary_dis)/(2*a)
            x2 = (-b - imaginary_dis)/(2*a)
            return (x1, x2)
    
    else:
        return (0j, 0j)
```

<!-- grade_id: Q7B, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(roots(1,2,3), ((-1+1.4142135623730951j), (-1-1.4142135623730951j)))
assert_equal(roots(0,0,4), (0j, 0j))
assert_equal(roots(-4,3,2), (-0.42539052967910607, 1.175390529679106))
```

## 8. The Game of Nims

In this game, two players sit in front of a pile of 100 stones. They take turns, each removing between 1 and 5 stones (assuming there are at least 5 stones left in the pile). The person who removes the last stone(s) wins.

In this problem, you’ll write a function to play this game. It may seem tricky, so break it down into parts. Like many programs, we have to use nested loops (one loop inside another). In the outermost loop, we want to keep playing until we are out of stones. Inside that, we want to keep alternating players. You have the option of either writing two blocks of code, or keeping a variable that tracks the current player.

We are going to use the `input` function to obtain the user's move. We might want to have an innermost loop that checks if the user’s input is valid. Is it a number? Is it a valid number (e.g. between 1 and 5)? Are there enough stones in the pile to take off this many? If any of these answers are no, we should tell the user and re-ask them the question.

If you choose to write two blocks of code, the basic outline of the program should be something like this:

    while [pile is not empty]:
       while [player 1’s answer is not valid]:
          [ask player 1]
       [execute player 1’s move]
       [same as above for player 2]
       
Be careful with the validity checks. Specifically, we want to keep asking player 1 for their choice as long as their answer is not valid, BUT we want to make sure we ask them at least ONCE. So, for example, we will want to keep a variable that tracks whether their answer is valid, and set it to `False` initially.

There is no need to limit yourself to only one function! Use as many functions as you want in order to make the code more readable and modular.

When you're finished, test each other's programs by playing them!

<!-- grade_id: Q8A, points: 4, grade: true, solution: true -->
```python
def is_valid(player_input, pile, max_stones):
    '''
    function to check whether input is valid. Will be called by player_input. Returns a boolean.
    '''
    return player_input <= max_stones and player_input >= 1 and player_input <= pile

def player_input(player_num, pile, max_stones):
    '''
    function that prompts user, calls is_valid, then returns the number of stones the player wants to remove
    '''
    stones = int(input("Player " + str(player_num) + ": You can remove up to " + str(max_stones) + " stones if available. How many?"))
    while not is_valid(stones, pile, max_stones):
        stones = int(input("Player " + str(player_num) + ": You can remove up to " + str(max_stones) + " stones if available. How many?"))
    return stones

def play_nims(pile, max_stones):
    '''
    An interactive two-person game; also known as Stones.
    @param pile: the number of stones in the pile to start
    @param max_stones: the maximum number of stones you can take in one turn
    Calls the function player_input
    '''

    while pile > 0:

        player1_stones = player_input(1, pile, max_stones)
        
        pile = pile - player1_stones

        if pile == 0:
            print("Player 1 wins")
            break

        else:
            player2_stones = player_input(2, pile, max_stones)
            
            pile = pile - player2_stones
            
            if pile == 0:
                print("Player 2 wins")
                break

    print("Game over")
```